### Matemātika programmētajiem. Programmēšanas valoda Python.

# 3. praktiskais darbs - Saraksti

1. Sarakstu (`list`) deklarē izmantojot kvadrātiekavas un atdalot elementus ar komatu. Sarakstam nav noteikts elementu skaits - to var patvaļīgi izvēlēties. Elementi var būt dažādu tipu.<br>
Ar funkcijas `len(...)` palīdzību iespējams noskaidrot elementu skaitu sarakstā.<br>
Izmēģini šos piemērus: <br>
<br>
len( [] ) <br>
len( [0, 1, 2, 3, 4] ) <br>
len( ["a", 1, [] ] ) <br>

```python
len( [] )
```

2. Izmantojot `append`, var saraksta beigās pievienot jaunu elementu.

```python
x = []
x.append(1)
x.append(2)
print(x)
```

3. Ja nepieciešams pievienot sarakstam vairākus elementus uzreiz, to var izdarīt ar `extend`  <br>
**Piezīme**: tā kā Jupyter Notebook šūnas "redz" viena otru, nav nepieciešams katru reizi veidot jaunu sarakstu - var turpināt darbu ar jau esošo

```python
x.extend([3, 4])
print(x)
```

4. Divus sarakstus var saskaitīt kopā

```python
y = [5, 6] + [7, 8]
print("y = ", y)
```

5. Sarakstus var reizināt ar veselu skaitli

```python
y=["a"] * 3
print("y = ", y)
```

6. Ir jābūt uzmanīgiem piešķirot sarakstam vērtības - piešķirot vienam sarakstam cita saraksta vērtības, saraksts netiek kopēts, bet gan tiek veidota norāde uz oriģinālo sarakstu. Veicot izmaiņas vienā sarakstā, tās pašas izmaiņas tiks veiktas arī otrā.

```python
print("x :", x)
y = x
x.append(5)
print("x :", x)
print("y :", y)
x == y
```

Ja ir nepieciešami divi saraksti ar vienādām vērtībam, bet bez norādes uz oriģinālo sarakstu (lai veicot izmaiņas vienā, tās neatspoguļojas otrā), tad jāizmanto `copy`

```python
y = x.copy()
x.append(6)
x != y
```

Ir **ļoti svarīgi** saprast un iegaumēt atšķirību starp piešķiršanas un kopēšanas operatoriem!

6. Līdzīgi kā tas bija `string`, ir iespējams pārbaudīt, vai sarakstā ir kāda konkrēta vērtība.

```python
a = [1, 2, 1, 2, 3, 3, 5]
4 in a
## BEGIN SOLUTION
2 in a
## END SOLUTION
```

Izmainiet doto kodu, lai tas parbaudītu vai saranstā `a` ir vērtība 2

7. Izmantojot `count(...)` var noskaidrot, cik reizes kāda vērtība atkārtojas sarakstā. Piemēra tiek atgriezts, cik reizes vērtība $3$ atkārtojas sarakstā `a`.<br>
Noskaidrojiet, cik reizes atkārtojas vērtība $1$ un cik reizes atkārtojas vērtība $7$

```python
a.count(3)
## BEGIN SOLUTION
print("repeats of 1 :", a.count(1))
print("repeats of 7 :", a.count(7))
## END SOLUTION
```

8. Elementu secība sarakstā nemainās, ja vien tā netiek speciāli mainīta. Pateicoties tam, mēs varam izmantot elementu **indeksus** - veselus skaitļus, kas apzīmē elementa atrašanās vietu. <br>
Programmēšanas valodās prasti sāk skaitīt no **nulles** - pirmajam  elementam būs indekss 0, otrajam elementam būs indekss 2, utt.

```python
a[0]
## BEGIN SOLUTION
a[3]
## END SOLUTION
```

Izmaniet kodu, lai noskaidrotu, kurš elements atrodas saraksta `a` vietā ar indeksu $3$ 

9. Strādājot ar ptogrammēšanas valodu Python, ir iespējams piekļūt elementiem arī izmantojot negatīvus skaitļus, piemēram pēdējā saraksta elementa indekss ir -1. <br>
Tas ir noderīgi, jo, nav nepieciešams zināt saraksta elementu skaitu, lai aprēķinātu pareizo indeksu:

```python
a[-1]
a[len(a)-1] # not neccecary, because we can do the same thing easier(see previous line)
```

10. Kāda vērtība atrodas saraksta priekšpēdējā indeksā?

```python
## BEGIN SOLUTION
a[-2]
## END SOLUTION
```

11. Ja zināma elementa vērtība, to masīva var izdzēst izmantojot `remove(...)`

```python
b = [1, 5, 9, 2, 1, 5, 10, -4, 5, 1, 6, 3]
b.remove(5)
## BEGIN SOLUTION
b.remove(1)
## END SOLUTION
print(b)

```

Kas notiek, ja vēlas izdzēst elementu pēc vērtības, bet šāda vērtība sarakstā atkārtojas vairākas reizes?

11. Ja zināms elementa indekss, to no saraksta var izdzēst lietojot atslēgasvārdu **del**

```python
del b[1]
```

12. Sarakstā var piekļūt ne tikai atsevišķiem elementiem, bet arī veselai elementu apakškopai. Programmēšanas valoda Pyhton to sauc par **segmentu** (slice). <br>
Lai piekļutu segmentam, ir jānosaka tā sākuma indekss(`slice_start`) un beigu indekss(`slice_end`). 

```python
c = [1, 5, 9, 2, 1, 5, 10, -4, 5, -3, 6, 3]
slice_start=3
slice_end=7
```

Izmantojot šos divus indeksus, varam piekļūt trim dažādiem segmentiem:
- segmentam, kas sākas ar `slice_start` (ieskaitot) un beidzas saraksta beigās
- segmentam, kas sākas saraksta sākumā un beidzas ar `slice_end`(neieskaitot)
- segmentam, kas sākas ar `slice_start` (ieskaitot) un beidzas ar `slice_end`(neieskaitot)

```python
c[slice_start : ] # from slice_stat(inclusive) to end of list
```

```python
c[ : slice_end] # from beginning of the list to slice_end(exclusive)
```

```python
c[slice_start : slice_end] # from slice_start (inclusive) to slice_end (exclusive)
```

Veicot izmaiņas dotajā kodā, izmēģiniet segmentu veidošanu ar citiem piemēriem, lai labāk izprastu, kā tā darbojas.

**Piezīme**: Ievērojiet, ka sākuma indekss vienmēr tiek iekļauts atslēgumā, bet beigu indekss netiek iekļauts. Ja kā kā sākuma un beigu indeksu uzstāda kā vienu un to pašu vērtību, tas sadalīs sarakstu divos sarakstos, kas nepārklāsies. Tas ir nodemonstrēts nākošajā koda šūnā.

```python
c == c[ : 6] + c[6 : ]
```

13. Elementus sarakstā ir iespējams sakārtot. Lai to izdarītu, elementiem jābūt **sakārtojamiem** - tadiem, ko vērtības iespējams salīdzināt izmantojot salīzināšanas operatorus ($==$, $<$, $<=$, $>$, $>=$) <br>
Ir divi iespējamie kārtošanas kārtošanas paņēmieni. Izmantojot `sort(...)`, elementi tiek sakārtoti tajā pašā sarakstā, kur tie atrodas. 

```python
print(c)
c.sort()
print("c after using sort:", c)
```

Izmantojot `sorted(...)`, teik izveidota saraksta kopija, kura tiek sakārtota, bet oriģinālais saraksta paliek bez izmaiņām.

```python
d=[1, 5, 9, 2, 1, 5, 10, -4, 5, -3, 6, 3]
e=sorted(d)
print("d after using sorted:", d)
print("e:", e)

```

`sort` un `sorted` sakārto elementus augošā secībā. Ja nepieciešams sakārtot elementus dilstošā jeb pretējā secībā (`reverse`), tad to norāda kā otro argumentu

```python
sorted([-5, 5, 0, 9, 19, -1], reverse = True)
```

14. Lai atrastu mazāko saraksta elementu, ir paredzēta funkcija `min(...)`. Lielākā elementa atrašanai ir funkcija `max(...)`.<br>
Atrodiet saraksta `d` lielāko un mazāko elementu

```python
## BEGIN SOLUTION
min(d)
max(d)
## END SOLUTION
```

15. Uzrakstiet savu funkciju `my_max`, kas, izmantojot `sort` vai `sorted`, atrod lielāko elementu sarakstā.

```python
## BEGIN SOLUTION
def my_max(m) :
    n=sorted(d, reverse=True)
    return n[0]
print(my_max(d))
## END SOLUTION

```

16. Ir iespējams iegūt sarakstu ar elementiem pretēja seībā, nekā dotā saraksta elementi. Piemēram, no saraksta [1, 9,-8, 0] iegūt sarakstu [0, -8, 9, 1]

```python
l1=[1, 9,-8, 0] 
l2=l1[::-1]
print(l2)
```

Kura kārtošanas funkcija ir piemērotāka šaja situācijā? Kāpēc?

17. Funkcija `sum(...)` atgriež visu saraksta elementu summu.<br>
Izlabo kodu, lai funkcija `mean` atgrieztu saraksta aritmētisko vidējo.<br>
**Padoms**: Aritmētiskais vidējais tiek aprēķināts pēc firmulas $a_1$; $a_2$;$...$; $a_n$ is $ \frac{1}{n} \sum_{i=1}^{n}=a_i$

```python
def mean(z) :
    return sum(z)

## BEGIN SOLUTION
def mean(z) :
    return sum(z)/len(z)
## END SOLUTION
print(mean(d))
```

## Izejas biļete

i1 Uzrakstiet funkciju `is_empty`, kas, gadījumā, ja saraksts ir tukšs, atgriež `True`, un pretēja gadījumā atgriež `False`

<!-- grade_id: i1_solution, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
def is_empty(a) :
    if len(a)==0 :
        return True
    else :
        return False

## END SOLUTION
```

<!-- grade_id: i1_test, points: 2, grade: true, solution: false -->
```python
## BEGIN HIDDEN TESTS
a=[]
b=[1, 5, 9]
assert is_empty(a)==True
assert is_empty(b)==False
## END HIDDEN TESTS
```

i2 Uzrakstiet skriptu, kas lūdz lietotājam ievadīt 10 skatiļus, saglabā tos sarakstā un izvada cik no tiem ir:
- pozitīvi skaitļi
- pāra skatiļi

<!-- grade_id: i2_solution, points: 3, grade: true, solution: true -->
```python
## BEGIN SOLUTION
pos=0
even=0
def make_l(a, pos, even) :
    if len(a)!=10 :
        tmp=int(input("please write a number: "))
        a.append(tmp)
        if tmp%2==0 :
            even+= 1
        if tmp>0 :
            pos+= 1
        make_l(a, pos, even)
    else :
        print("There are ", pos, "positive numbers and ", even, " even numbers on the list")

b=[]
make_l(b, pos, even)

## END SOLUTION
```

i3 Uzrakstiet skriptu, kas izvada elementus, kas kopīgi abiem sarakstiem (atrodas gan saranstā `a`, gan sarakstā `b`)

<!-- grade_id: i3_answer, points: 1, grade: true, solution: true -->
```python
a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89]
b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]
## BEGIN SOLUTION
new_list = []

for num in a:
    if num in b:
        new_list.append(num)

print(new_list)
## END SOLUTION
```

i4 Uzrakstiet skriptu, kas lūdz lietotājam ievadīt vārdu un pārbauda vai šis vārds ir palindroms (tāds vārds, kas ir vienāds lasot abos virzienots, piemēram "anna"
**Padoms**: Ir iespējams iegūt virkni pretējā secībā, izmantojot to pašu paņēmienu, kādu izmantojām, lai iegūtu sarakstu pretējā secībā

<!-- grade_id: i4_answer, points: 2, grade: true, solution: true -->
```python
## BEGIN SOLUTION
word=input("Please enter a word: ")
word=str(word)
reverse=word[::-1]
print(reverse)
if word == reverse:
    print("This word is a palindrome")
else:
    print("This word is not a palindrome")
## END SOLUTION
```