### Matemātika programmētajiem. Programmēšanas valoda Python.

# 4. praktiskais darbs - Atkārtojuma konstrukcijas

Viena no lielākajām datora priekšrocībam ir tā spēja nenogusrstot atkārtot vienas un tās pašas darbības daudz reizes. Konstrukcijas, kas ļauj izpildīt koda fragmentu nepieciešamo reižu skaitu sauc par **atkārtojuma konstrukcijām** jeb **cikliem**. Katru cikla atkārtojumu sauc par **iterāciju**<br>
Programmēšanas valodā Python ir divu veidu cikli - **while**("kamēr") cikls jeb cikls ar nosacījumu un **for**("priekš") cikls jeb cikls ar skaitītāju. <br>

### `for` cikls

1. `for` ciklu Python valodā var uztvert kā "priekš katra elementa secībā" (for each element in sequence). Cikla ķermenis tiks izpildīts vienu reizi katra iterācijā.<br>

```python
numbers = [0,1,2,3,4]

for i in numbers :
    print(i)
```

2. Cikli, kam jāizpildās noteiktu reižu skaitu, ir tik bieži sastopami, ka Python ir pieejama īpaši tam paredzēta intervāla funkcija - `range(...)`

```python
for i in range(5) :
    print(i)
```

Ja `range` kā parametrus dod divus skaitļus, tad tas darbojas līdzīgi ka iepriekš apskatītais `slice` - nosakot sākumu un beigas intervala, kurā notiks iterācijas. Ta pat ka `slice`, intervāla pirmais elemets tiek iekļauts, bet pēdējais netiek iekļauts.<br>
Iespējams pievienot arī trešo mainīgo - tas noteiks "soļa" garumu. Dots piemērs, kas intervalā no 3 (ieskaitot) līdz 11 (neieskaitot) izvada katru otro skaitli

```python
for i in range(3, 11, 2) :
    print(i)
```

Izmēģiniet iepriekšējos uzdevumos apskatītos piemērus ar savām vērtībam, lai pārliecinātos, ka sparotat kā strādā `for` cikli.

3. Uzrakstiet kodu, kas izvadīs šādu no zvaigznītēm veidotu trijstūri:
<br>
$*$ <br>
$**$ <br>
$***$ <br>
$****$ <br>
$*****$ <br>

```python
## BEGIN SOLUTION
n=6;
for i in range(n):
    for j in range(i):
        print ('* ', end="")
    print('')
## END SOLUTION
```

4. Izmantojot ciklus darbībam ar masīviem, nepieciešams asociēt $i$-to elementu masīvā ar indeksu $i$. Vienkārši iterēšana cauri masīvam ar ciklu, nedod indeksus. Lai atrisinātu šo problēmu, varam izmantot funkciju `len` kā parametru funkcijai `range`

```python
x = ["a", "b", "c", "d", "e", "f"]

for i in range(len(x)) :
    print(i, x[i])
## BEGIN SOLUTION
for i in range(len(x)) :
    if i%2 !=0 :
        print(i, x[i])
## END SOLUTION
```

Veiciet izmaiņas dotajā kodā, lai tiktu izvadīti tikai elementi, kas masīvā atrodas nepāra indeksos.

5. Dažreiz ir nepieciešams iterēt cauri elementiem tikai līdz brīdim, kad izpildās kāds nosacījums un tad iziet no cikla. Šajā gadījumā izmantojam `break`

```python
for i in "story time always goes on" :
    if i == "a" :
        break
    print(i)

print("has ended")
```

Iespējams arī, ka nepieciešams iterēt cauri visiem elementiem, bet gadījumā, ja elements šajā iterācijā atbilst kādam konkrētam nosacījumam, tad to izlaist un turpināt ar nākošo elementu. Šajā gadījumā izmantojam `continue`

```python
for i in range(15) :
    if i%3==0 :
        continue
    print(i)
    
## BEGIN SOLUTION
for i in range(15) :
    if i%3==0 or i%2==0:
        continue
    print(i)
## END SOLUTION
```

Šajā piemērā tiek izvadīti skaitļi no 1 līdz 15, izlaižot skaitļus, kas bez atlikuma dalās ar 3. Izmainiet kodu, lai tiktu izlaisti skaitļi, kas dalās bez atlikuma ar 3 vai 2 <br>
**Padoms**: izmantojiet `or`

6. Iepriekšējā praktisjakā darbā tikām rakstījuši funciju `mean`, kas aprēķina aprēķina aritmētisko vidējo izmantojot funkciju `sum` un funkciju `len`. Izmainiet kodu, lai `sum` un 'len' vietā elementu summas un elementu skaita tiktu izmantots `for` cikls.

```python
def mean(z) :
    return sum(z)/len(z)

## BEGIN SOLUTION
def mean(z) :
    s=0
    count=0
    for i in range(len(z)) :
        s +=z[i]
        count +=1
    return s/count       
## END SOLUTION

d=[1, 5, 9, 2, 1, 5, 10, -4, 5, -3, 6, 3]
print(mean(d))
```

7. Definējiet funkciju `normalise`, kam tiek padots skaitļu saraksts un kas atgriež sarakstu, kur katra tā vērtība ir  normalizēta - visu elementu summa ir 1.0 un to relatīvais sadalījums ir tāds pats, kā elementiem ievaddatu sarakstā.<br>
Piemēram, padodot sarakstu [4, 8, 4] tiktu atgriezts saraksts [0.25, 0.5, 0.25]

```python
def normalise(z) :
    
    ## BEGIN SOLUTION
    norm_list=z.copy()
    for i in range(len(z)) :
        norm_list[i]=z[i]/sum(z)
    return norm_list    
    ## END SOLUTION
```

8. Programmēšanas valodā Python ir iespējams veidot testa piemērus izmantojot atslēgasvardu `assert`. Testa piemēri paredzēti, lai pārbaudītu, vai iepriekš uzrakstītais kods darbojas pareizi. Lai uzrakstītu testa piemēru, nepieciešams zināt kādi būtu pareizie izvaddati pie konkrētiem ievaddatiem.<br> 

```python
assert normalise([12, 3, 6, 3])==([0.5, 0.125, 0.25, 0.125])
```

Ja šūna ar testa piemēriem atgriež kļūdas paziņojumu, tas nozīmē, ka kods, kas tiek testēts, nestrādā pareizi. <br>
Pat ja netiek atgriezs kļūdas paziņojums, tas negarantē, ka kods ir pareizs - var gadīties, ka kods strādā pareizi tikai pie konkrētajām vērtībām, tāpēc ieteicams vienmēr rakstī vairākus testa piemērus. <br>
Uzrakstiet vēl vismaz divus testa piemērus funkcijai `normalize`

### `while` cikls

9. `while` cikla ķermenī norādītās darbības veiks tik ilgi, kamēr izpildīsies  noteikts nosacījums.

```python
count = 0

while count < 5:
    print("count =", count)
    count += 1
    
## BEGIN SOLUTION
count=5    
while count >= 5 and count <=10:
    print(count)
    count += 1
## END SOLUTION
```

Izmainiet kodu, lai `while` cikls izvadītu skaitļus no 5 līdz 10

10. Gadījumos, kad `while` cikla nosacījums ir ļoti garš, koda pārskatāības nolūkos mēdz veidot bezgalīgu ciklu un izmantot `break` iziešanai no tā.  Izlabojiet doto kodu, lai tā būtu bezgalīga cilpa, no kuras darbība tiek pārtraukta gadījumā, ja lietotājs ievada `y`

```python
attempts = 0

while True :
    response = input("Do you want to quit? (y/n): ")
    attempts += 1
    if response == "y" :
        break

print("Exiting after", attempts, "attempts")
```

## Izejas biļete

i1 uzveidojiet funkciju `composite(...)`, kas  
- atrod visus dotā skaitļa dalītajus
- ja skaitlis dalās tikai pats ar sevi un ar 1, atgriež, ka tas ir pirmskaitlis - `prime`
- ja tas nav pirmskaitlis, atgriež visus tā dalītājus

<!-- grade_id: i1_solution, points: 0, grade: false, solution: true -->
```python
def composite(a) :

    ## BEGIN SOLUTION
    m = []  
    for i in range(1, a) :
        if a%i == 0 :
            m.append(i)
    
    if len(m)<=2 :
        return "prime"
    else :
        return m
    ## END SOLUTION   
```

<!-- grade_id: i1_test1, points: 2, grade: true, solution: false -->
```python
assert composite(12) == [1, 2, 3, 4, 6]
```

<!-- grade_id: i2_test2, points: 2, grade: true, solution: false -->
```python
assert composite(5) == "prime"
```

i2 Uzrakstiet skriptu, kas lūdz lietotājam ievadīt skaitli. Ja ievadītais skaitlis ir lielāks par 10, izvada paziņojumu, ka ievadīts pareizs skaitlis. Ja ievadītais skaitlis ir mazāks par 10, tad lūdz ievadīt skaitli atkārtoti, līdz tiek ievadīts skaitlis, kas ir lielāks par 10.

<!-- grade_id: i2_solution, points: 2, grade: true, solution: true -->
```python
## BEGIN SOLUTION
count=5    
while True :
    response = int(input("Please write a number that is larger than 10: "))
    if response > 10 :
        print("input is correct")
        break
    else: 
        print("input is not correct, please try again")
## END SOLUTION
```

i3 Grāmatu sērijā ir 6 grāmatas, katra no tām maksā 8 € <br>
Ja pērk 2 grāmatas no sērijas, tiek piešķirta 5% atlaide un ar katru nākošo grāmatu atlaide pieaug par 5%. Piemēram, pērkot  grāmatas, tiks piemērota 15% atlaide. Uzrakstiet funkciju `get_price`, kam tiek padots grāmatu skaits un kas atgriež cik par tām jāmaksā.

<!-- grade_id: i3_solution, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
def get_price(num) :
    price = 0
    discount = 0
    if num>6 or num<1 :
        print("invalid book count")
    elif num == 1 :
        return 8
    else :
        for i in range(num) :
            price += 8
            if num >=2 :
                discount +=5
        return price -(price*(discount*0.01))
            
print(get_price(6))        
## END SOLUTION
```

i4 Uzrakstiet testa piemēru funkcijai `get_price`, kas pārbauda, vai pie argumenta 6 tiek atgriezta vērtība 33.6

<!-- grade_id: i4_solution, points: 1, grade: true, solution: true -->
```python
## BEGIN SOLUTION
assert get_price(6)==33.6
## END SOLUTION
```

i5 Uzrakstiet skriptu, kas lūdz lietotājam ievadīt personu skaitu (pozitīvu veselu skaitli) `person_count` un padot šo skaitli funkcijai `make_list_of_ages`.
Izveidojiet funkciju `make_list_of_ages`, kas saņem kā parametru lietotāja ievadīto skaitli un izmantojot `for` ciklu aiazpilda saraktu
- sarakta vērtību skaitu nosaka argument, ko saņem funkcija
- katrā iterācijā lietotājam lūdz ievadīt dzimšanas gadu un pievieno sarakstam, cik šobrīd būtu gadu personai, kas dzimusi tajā gadā (2019-`year_of_birth`)
- ja ievadīts gads, kas vēl nav pienācis, kā vecums tiem ierakstīts 0

<!-- grade_id: i5_answer, points: 3, grade: true, solution: true -->
```python
## BEGIN SOLUTION
def make_list_of_ages(num) :
    list_of_ages=[]
    for i in range(num) :
        year_of_birth = int(input("Please input year of birth: "))
        if year_of_birth>2019 :
            list_of_ages.append(0)
        else: 
            list_of_ages.append(2019-year_of_birth)
    return list_of_ages
        
        
person_count=int(input("How many persons do you want to add? "))
if person_count<=0 :
    print("Number not valid")
else: 
    print(make_list_of_ages(person_count))
    
## END SOLUTION
```