### Matemātika programmētajiem. Programmēšanas valoda Python.

# 6. Praktiskai darbs - Divdimensiju matricas un lineāru vienādojumu sistēmas

### Divdimensiju matricas

1.Programmēšanas valodā Python divdimensiju matrica var attēlot kā sarakstu, kas sastāv no sarakstiem - katra matricas rinda tiek uztverta kā atsevišķs saraksts. <br>
Lai izprastu, kā tas notiek, apskatīsim matricas $M=\begin{pmatrix} 
9 & 7 & 6\\
5 & 1 & 0
\end{pmatrix}$ attēlojumu

```python
m = ([9, 7, 6], [5, 1, 0])
print(m)
```

Šobrīd izvadītie skaitļi ne pārāk izskatās pēc matricas. Izmantosim funkciju `matrix` no moduļa `numpy`, lai "pateiktu", ka vēlamies izvadīt sarakstu kā matricu

```python
import numpy
print(numpy.matrix(m))
```

Izveidojiet savu funkciju, kas izsauc šo `numpy` funkciju un izvada matricu

```python
def printMatrix(s):
    
## BEGIN SOLUTION
    print(numpy.matrix(s))
## END SOLUTIN
printMatrix(m)
```

2. Izveidosim aizplidītu ar nullēm matricu, kam ir 4 rindas un 5 kolonnas

```python
rows, columns = 4, 5
arr = [[0 for i in range(columns)] for j in range(rows)]
printMatrix(arr)
```

3. Lai piekļūtu matricas elementam, ir jāzina tā atrašanās vieta matricā - kurā rindā un kurā kolonnā tas atrodas. <br>
Ievietosim tikko izveidotā masīva 2. rindas 4. kolonna vērtību 9

```python
arr[1][3]=9
printMatrix(arr)
```

Elements tika ievietots 2.rindas 4.kolonnā, bet skriptā tika norādīts [1][3] <br>
Kāpēc?

Ievietojiet elementu ar vērtību 5 3.rindas 4.kolonnā 

```python
## BEGIN SOLUTION
arr[2][1]=5
printMatrix(arr)
## END SOLUTION 
```

4. Izmantosim `for` ciklus, lai aizpildītu visus matricas elementus ar skaitļiem augošā secībā

```python
n=0
for i in range(rows):
    for j in range(columns):
        arr[i][j]=n
        n+=1
printMatrix(arr)

## BEGIN SOLUTIN
for i in range(rows):
    for j in range(columns):
        if j%2==0 :
            arr[i][j]=n
            n+=1
printMatrix(arr)
## END SOLUTION
```

Izpētiet tikko izmantoto kodu un izmainiet to tā, lai tiktu aizpildīta tikai katra otrā kolonna

5.Saskaitot divas matricas, saskaita to atbilstošos elementus.
**Atgādinājums**: Saskaitīt un atņemt drīkst tikai vienāda izmēra matricas!

```python
arr1=([1, 2, 3], [1, 2, 3], [1, 2, 3])
arr2=([1, 0, 0], [0, 1, 0], [0, 0, 1])

arrSum = [[0 for i in range(3)] for j in range(3)]

for i in range(3):
    for j in range(3):
        arrSum[i][j] = arr1[i][j] + arr2[i][j]
        
printMatrix(arrSum)
```

6. Veicot matricu atņemšanu, atņem to atbilstošos elementus. Izmantojot iepriekšējo piemēru kā paraugu, uzrakstiet matricu atņemšanu

```python
##BEGIN SOLUTION
arrMin = [[0 for i in range(3)] for j in range(3)]
for i in range(3):
    for j in range(3):
        arrMin[i][j] = arr1[i][j] - arr2[i][j]
        
printMatrix(arrMin)
##END SOLUTION
```

7. Matricas reizināšana ar veselu skaitli notiek ļoti līdzīgi - katrs matricas elements jāsareizina ar doto skaitli. Uzrakstiet skriptu matricas `arr1` reiznāšanai ar 7

```python
##BEGIN SOLUTION
arrMultiplByNumber = [[0 for i in range(3)] for j in range(3)]
for i in range(3):
    for j in range(3):
        arrMultiplByNumber[i][j] = arr1[i][j] * 3
        
printMatrix(arrMultiplByNumber)
##END SOLUTION
```

9. Šīs darbības ir iespējams veikt arī ar moduļa `numpy` palīdzību. <br>
Piemērā parādīta saskaitīšana izmantojot `numpy` moduli

```python
arrSum= numpy.array(arr1)+numpy.array(arr2)
printMatrix(arrSum)
```

Pēc šī piemēra, uzrakstiet matricu atņemšanu

```python
arrMin= numpy.array(arr1)-numpy.array(arr2)
printMatrix(arrMin)
```

10. Matricas reizināšana ar veslu skaitli notiek līdzīgi kā matricu saskaitīšana, tikai `+` aizstājot ar `*` un otrās matricas vieta ierakstot skaitli, ar ko vēlaties reizināt. <br>
Uzrakstiet skriptu matricas `arr1` reizināšanai ar 8

```python
## BEGIN SOLUTION
arrMulti= numpy.array(arr1)*8
printMatrix(arrMulti)
## END SOLUTION
```

**Ieteikums**: Salīdziniet ar funkciju `dot` iegūto rezultātu un iepriekš uzrakstītās `matrixmultipl` rezultātu, lai pārbaudītu, vai `matrixmultipl` ir uzrakstīta pareizi.

<br>
11. Reizinājuma rezultātā iegūtās matricas rindu skaits ir tāds pats kā matricas A rindu skaits un kolonnu skaits sakrīt ar kolonnas B rindu skaitu. <br>
<br>
Reizinot iegūtajā matricā katrs [i][j] elements ir iegūts saskaitot visus matricas B j-tāš rindas atbilstošos reizinājumus<br> 
<br>
**Atgādinājums**: Matricas  A  un  B  reizināt drīkst tikai tad, ja matricas  A  kolonnu skaits sakrīt ar matricas  B  rindu skaitu. <br>
<br>
Tālāk ir redzama funkcija matricu reizināšanai, bet tajā ir kļūda. Atrodiet un izlabojiet to.

```python
a1=([0, 0, 0], [5, 5, 5], [8, 8, 8], [9, 9, 9])
b1=([1, 1, 1, 1], [2, 2, 2, 2], [4, 4, 4, 4])

def matrixmultipl (A, B) :
    C = [[0 for row in range(len(A))] for col in range(len(B[0]))]
    for i in range(len(A)):
        for j in range(len(B[0])):
            for k in range(len(B)):
                C[j][i] += A[i][k]*B[k][j]
    return C

printMatrix(matrixmultipl(a1, b1))
## BEGIN SOLUTIN
#it scould be C[i][j]
## END SOLUTION
```

Matricu reizināšanai var izmantot moduļa `numpy` funkciju `dot` 

```python
c=numpy.dot(a1, b1)
printMatrix(c)
```

12. Kāpinot matrica tiek reizināta pati ar sevi, tāpēc kāpināt drīkst tikai kvadrātiskas matricas

```python
a3=([6, 1, 1],[4, -2, 5],[2, 8, 7])
def powMatrix (ma, n) :
    if len(ma)!=len(ma[0]) :
        print("Operation not opssible, matrix is not square")
    else :
        for i in range(n) :
            res = numpy.dot(ma, ma)
    return res

c3=powMatrix(a3,2)
printMatrix(c3)

## BEGIN SOLUTIN
#it scould be C[i][j]
## END SOLUTION
```

13. Lai iegūtu matricas determinantu, var izmantot funkciju `linalg.det(...)` no moduļa `numpy`


```python
print(numpy.linalg.det(a3))
```

14. Izmantojot funkciju `help`, apskatiet `numpy.linag` dokumentaciju un atrodiet, ar kādu funkciju iespējams atrast inverso matricu.

```python
## BEGIN SOLUTION
help(numpy.linalg)
## END SOLUTION
```

Izmantojot atrasto funkciju, iegūstiet dotās matricas inverso matricu

```python
a5=([1, 2, 3], [0, 1, 5], [5, 6, 0])
## BEGIN SOLUTION
c4=numpy.linalg.inv(a5)
printMatrix(c4)
## END SOLUTION

```

15. Uzrakstiet skriptu, kas atgriež matricas `a5` transponēto matricu. <br>
**Atgādinājums** : transponēto matricu iegūst dotās matricas rindas elementus mainot vietām ar atbilstošās kolonnas elementiem

```python
##BEGIN SOLUTION
tmp=len(a5)
arrTransp = [[0 for i in range(tmp)] for j in range(tmp)]
for i in range(tmp):
    for j in range(tmp):
        arrTransp[i][j] = a5[j][i]
        
printMatrix(arrTransp)
##END SOLUTION
```

Transponēto matricu var iegūt ar moduļa `numpy` funkciju `matrix.transpose(...)`. Izmantojot to, iegūstiet matricas `a5` transponeto matricu lai pārbaudītu, vai iepriekšējais uzdevums izpildīts pareizi.

```python
## BEGIN SOLUTION
printMatrix((numpy.array(a5)).transpose())
## END SOLUTION
```

**Padoms**: Ja kāda fukcija sagādā grūtības, apskatiet tās dokumentāciju un pievērsiet uzmanību tur atrodamajiem piemēriem

### Lineāru vienādojumu sistēmas

16. Ir dota lineāra vienādojumu sistēma 
\begin{cases} 5x+3y-z=-1 \\ x-y+z=0 \\ 2x-3y-z=5  \end{cases}
Lai to atrisinātu izmantojot matricas un programmēšanas valodas Python piedāvātās iespējas, vispirms jāizveido sistēmas koeficentu matrica:

```python
a=([5, 3, -1], [1, -1, 1], [2, -3, -1])
```

Tad jāizveido brīvo locekļu matrica:

```python
b=[-1, 0, 5]
```

Tagad atrisināt sistēmu var izmantojot funkciju `linalg.solve` no moduļa `numpy`.

```python
res=numpy.linalg.solve(a, b)
printMatrix(res)
```

Iegūta matrica, kur pirmais elements ir pirmā vienādojuma atrisinājums, otrais ir otrā vienādojuma atrisinājums utt. 

Izmantojot `numpy.allclose(...)` iespējams pārbaudīt, vai kāds vērtību saraksts ir derīgs kā dotāš lineāras vienadojumu sistēmas atrisinajums. Piemēram, pārbaudīsim, vai [0.2, 2, 1.3] ir atrisinājums iepriekšējā uzdevumā dotajai sistēmai.

```python
x=[0.2, 2, 1.3]
numpy.allclose(numpy.dot(a, x), b)
```

Kā redzams, šie rezultāti neder kā dotās sistēmas atrisinājums.

## Izejas biļete

i1 Atrisiniet lineāro vienādojumu sistēmu
\begin{cases} 5x-2z=-1 \\ y+z=2 \\ 2x+y=6  \end{cases}

<!-- grade_id: i1_answer, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
a=([5, 0, -2],[0, 1, 1], [2, 1, 0])
b=[-1, 2, 6]
res=numpy.linalg.solve(a, b)
printMatrix(res)
##END SOLUTION

```

<!-- grade_id: i1_test, points: 2, grade: true, solution: false -->
```python
assert numpy.allclose(numpy.dot(a, res), b) == True
```

i2 Iegūstiet doto matricu reizinājumu

<!-- grade_id: i2_answer, points: 2, grade: true, solution: true -->
```python
a2=([1, 0, 5], [5, 8, 0], [4, 0, 1])
b2=([1, 5, 1, ], [3, 9, 2], [4, -5, -4])
## BEGIN SOLUTION
c2=numpy.dot(a2, b2)
printMatrix(c2)
## END SOLUTION

```